# Homework 1


## Problem 1 (If-statment)

If statement:

Given an integer n, perform the following conditional actions:

    If n is odd, set output to "Hello World"
    If n is even and in the inclusive range of 2 to 5, set output to "Goodbye World"
    If n is even and in the inclusive range of 6 to 20, set output to "Hello World"
    If n is even and greater than 20, set output to "Goodbye World"

Example:

    -- Input: n = 2
    -- Output: output = "Goodbye World"

```python

```

Write your code here

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
def problem1(n):
    output = 0
    
    if n%2 != 0 :
        output = "Hello World"
    else:
        if n>=2 and n<=5:
            output = "Goodbye World"
        elif n>=6 and n<=20:
            output = "Hello World"
        elif n>20:
            output = "Goodbye World"
    
    return(output)
```

<!-- grade_id: cell-2853176972493b6e, points: 25, grade: true, solution: false -->
```python
assert problem1(3) == "Hello World"
assert problem1(2) == "Goodbye World"
assert problem1(18) == "Hello World"
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 25, grade: true, solution: false -->
```python
assert problem1(29) == "Hello World"
assert problem1(24) == "Goodbye World"
```

Write ChatGPT code here

```python

```

```python

```

```python

```

## Problem 2 (For-loop, basic list manipulation)

You and $K-1$ friends want to buy $N$ different pieces of candy.  Candy piece $i$ costs $p_i$. The candy shop owner is greedy and wants to maximize his number of new customers, so he increases the sale price of the candy for repeated customers; more precisely, if a customer has already purchased $x$ pieces of candy, then the cost of candy $i$ is $C_i*(x+1)$.

Find and print the minimum cost for your group to purchase all $N$ pieces of candy.

The input:

    N: the number of candies to buy
    K: the number of people
    C: the cost of candy

The output:

    the total cost

**Hint:** You may want to use the list comprehension function.


Example:

    --Input: 
        N = 4
        K = 2
        C = [1,3,4,5]

    --Output:
        cost = 4 + 5 + 3 * 2 + 1 * 2 = 17

(Why 17 is the minimum cost in this case? You would always like your group to purchase the most expensive candies first. In this way, you will get the multipliers on candy costs for the cheap candies. In this case, two people in your group first purchase candies with costs 4 and 5. And then they will purchase candies with costs 3 and 1. When purchasing candies with costs 3 and 1, both of persons in your group are repeated customers, so they pay 3\*2 and 1\*2.)



```python

```

Write your code here:

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(N, K, C):
    #sort costs decreasing
    C = sorted(C,reverse = True)
    
    #variable to store the cost
    cost = 0
    
    # YOUR CODE HERE
    
    for i in range(N):
        if i % K == 0:
            times =  i // K
            cost += C[i] * (times) 
        else:    
            times = ( i // K ) + 1
            cost += C[i] * (times) 

    return(cost)

```

```python
def problem2(N, K, C):
    #sort costs decreasing
    C = sorted(C,reverse = True)
    
    #variable to store the cost
    cost = 0
    
    # YOUR CODE HERE
    for i in range(N):
        cost += C[i] * ((i // K) + 1)
        
    return(cost)
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 25, grade: true, solution: false -->
```python
# Test cases for Problem Set 4
import random
random.seed(100)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 8838
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 25, grade: true, solution: false -->
```python
import random
random.seed(10)
N=random.randint(1,100) 
K=random.randint(1,100)   
C = [random.randint(1,1000) for i in range(N)]
assert problem2(N, K, C) == 193878
```

```python

```

```python

```

```python

```

Write ChatGPT code here

```python
def minimum_cost(N, K, C):
    # Sort the candy prices in descending order
    C.sort(reverse=True)
    
    # Initialize the total cost
    cost = 0
    
    # Initialize the number of purchases for each person
    purchases = [0] * K
    
    # Iterate through the candies
    for i in range(N):
        # Find the person who will purchase the candy
        person_index = i % K
        
        # Calculate the cost for this candy for the person
        candy_cost = C[i] * (purchases[person_index] + 1)
        
        # Update the total cost and the number of purchases for the person
        cost += candy_cost
        purchases[person_index] += 1
    
    return cost

# Example usage:
N = 4
K = 2
C = [1, 3, 4, 5]

cost = minimum_cost(N, K, C)
print("Total cost:", cost)
```

```python

```

```python

```

```python

```

Please compare the difference of two codes above. What do you learn after using ChatGPT and how can improve your code? Please write below.

```python

```

```python

```