# Homework 2


## Problem 1 (Tuple):

Given a list of integers, create a tuple of those integers. Then compute and print the result of hash of the tuple.

Note: [hash()](https://docs.python.org/2/library/functions.html#hash) is one of the functions in the __builtins__ module, so it need not be imported.

[Here is the tutorial on why hash function is important in computer science.](https://www.hackerearth.com/practice/data-structures/hash-tables/basics-of-hash-tables/tutorial/)


Example:

    input: int_list = [1, 2, 3]
    output: 2528502973977326415

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
#Note: Based on your computer, your hash result could be different.  
import builtins

def problem1( input_list ):

    int_tuple = tuple(int_list)
    hash_result = hash(int_tuple)
            
    return(hash_result)


```

<!-- grade_id: cell-2853176972493b6e, points: 10, grade: true, solution: false -->
```python
assert problem1([1, 2, 3]) == 2528502973977326415
```

<!-- grade_id: cell-ab1b1d0a5f7b1fce, points: 10, grade: true, solution: false -->
```python
assert problem1([1, 2, 3, 4]) == 485696759010151909
```

Write ChatGPT Code here

```python

```

```python

```

## Problem 2 (Dictionary and String Methods):

You have a list of strings which contain student record. Each string contains the student's name, and their grades in Maths, Physics and Chemistry. The grades can be **floating values**. These values (names and grades) are separated by comma. You are required to save the record in a dictionary data type. The user then enters a student's name. Output the average percentage marks obtained by that student, correct to two decimal places.

Example:

    input:
        grades = ["Jake, 99, 70, 50", "Dennis, 100, 100, 98"]
        name = "Dennis"
    output: "The average grade of Dennis is 99.33"

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(grades, name):

    elements = {}
    
    for i in grades:
        
        line = i.split(", ")
        student = line[0]
        scores = line[1:]
        total = 0
        
        for i in scores:
            total += int(i)
        avg = total / len( line[1:] )
        
        elements[student] = round(avg, 2)
        
    output = "The average grade of " + name + " is " + str( elements[name] )

    return(output)

```

<!-- grade_id: cell-681343a4d7e22bbd, points: 10, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 54", "Dennis, 100, 100, 98"]
name = "Dennis"
assert problem2(grades, "Dennis") == "The average grade of Dennis is 99.33"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 10, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 55", "Dennis, 100, 100, 98"]
name = "Jake"
assert problem2(grades, "Jake") == "The average grade of Jake is 74.67"
```

Write ChatGPT Code here

```python

```

```python

```

## Problem 3 (List Methods):

You are given N positive numbers. Store them in a list and find the second largest number.
(If there are two largest numbers, the second largest number must be smaller than both of them)

Example:

    input: int_list = [1,3,1,4,4,2]
    ouput: 3

<!-- grade_id: cell-f16b41b82fe73df0, points: 0, grade: false, solution: true -->
```python
def problem3(int_list):
    
    second_largest = 0
    
    int_list = sorted(int_list)
        
    int_list = [ x for x in int_list if x != int_list[-1]]
    
    second_largest = [ x for x in int_list if x == int_list[-1]]
    
    result = second_largest[0]
    
    return(result)
```

<!-- grade_id: cell-b9a67e5a1cc6b7eb, points: 10, grade: true, solution: false -->
```python
assert problem3([1,3,1,4,4,2]) == 3
```

<!-- grade_id: cell-2e2644d4a46134fc, points: 10, grade: true, solution: false -->
```python
assert problem3([1,2,3]) == 2
```

Write ChatGPT Code here

```python

```

```python

```

## Problem 4 (All-in-one: list + string + file read):

We downloaded several stock's daily level data for you in ./data/hw2_data/stock_data/. You will write code that takes stock tickers as inputs and output one dictinoary. The dictionary uses the stock tickers as keys and the values are these stocks' average open prices in 2017 (until the end of the file).

Exmaple:

    input: tickers = ["AAPL"]
    output: {'AAPL': 143.465799260355}


<!-- grade_id: cell-0b99248c6f4c937e, points: 0, grade: false, solution: true -->
```python
def problem4(tickers):

    return_dict = {}
    
    for i in range( len(tickers) ):
        files = open("Homework-2/data/stock_data/%s.csv" %tickers[i])
        lines = files.readlines()
        name = str( tickers[i] )
        total = 0
        time = 0
        for i in range( len(lines) ):
            lines[i] = lines[i].split(",")
            if i != 0:
                total += float( lines[i][2] )
                time += 1
                avg = total / time
                return_dict[ name ] =  avg 

    return(return_dict)

problem4(["AAPL", "GOOG"])
```

Output:
```
{'AAPL': 143.465799260355, 'GOOG': 884.3826611893493}
```

<!-- grade_id: cell-59169f1e7e9bdbdc, points: 10, grade: true, solution: false -->
```python
assert problem4(["AAPL"]) == {'AAPL': 143.465799260355}
```

<!-- grade_id: cell-14e4a8c2e0fdedf6, points: 10, grade: true, solution: false -->
```python
assert problem4(["AAPL", "GOOG"]) == {'AAPL': 143.465799260355, 'GOOG': 884.3826611893493}
```

Write ChatGPT Code here

```python

```

```python

```

## Problem 5 (Power Set Problem) 

For a given list L, generate a list of lists which contain all subsets of L.  You are not allowed to use Python helper packages to do this.

Example:

Input: 

    L = [1,2]

Output:

    power_set = [[], [1], [2], [1,2]]

<!-- grade_id: cell-d3b79352a6e89063, points: 0, grade: false, solution: true -->
```python
def problem5(original_set):
    
    def backtrack( index, current_subset):
        power_set.append( current_subset[:] )
        
        for i in range( index, len( original_set ) ):
            current_subset.append( original_set[i] )
            backtrack( i + 1, current_subset )
            current_subset = []
                       
    power_set = []
    backtrack( 0, [] )
    return power_set

sorted(problem5([1]))

```

Output:
```
[[], [1]]
```

<!-- grade_id: cell-a2097c9f7e471e1c, points: 10, grade: true, solution: false -->
```python
assert sorted(problem5([1])) == sorted([[], [1]])
```

<!-- grade_id: cell-cfabbede7b5a2d5b, points: 10, grade: true, solution: false -->
```python
assert sorted(problem5([1, 2])) == sorted([[], [1], [2], [1,2]])
```

Write ChatGPT Code here

```python

```

```python

```

Please compare the difference of two codes above. What do you learn after using ChatGPT and how can improve your code? Please write below.

```python

```

```python

```