# Homework 3


## Problem 1 (Reocurring Characters: dictionaries) 

For a given string s, return the first reoccuring character in the string.  Try to get an approach that is $O(n)$, where $n$ is the number of characters in the string. If there is no such character, return None. 
    
    -- Input:"ABCA"
    -- Output:"A"

```python
def problem1(s):
    
    occurent_letter = None
    temp = set()
    
    for i in range( len(s) ):
        
        if s[i] in temp:
            occurent_letter = s[i]
            return(occurent_letter)
        else:
             temp.add(s[i])
                
    return None

```

<!-- grade_id: cell-2853176972493b6e, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCA") == "A"
```

<!-- grade_id: cell-b16dd3056b5a9f95, points: 10, grade: true, solution: false -->
```python
assert problem1("ABC") == None
```

<!-- grade_id: cell-ab1b1d0a5f7b1fce, points: 10, grade: true, solution: false -->
```python
assert problem1("ABCBA") == "B"
```

<!-- grade_id: cell-b78086c290d4a60c, points: 0, grade: false, solution: true -->
```python
#Write ChatGPT code here

```

## Problem 2 (String Methods):

We have learnt that strings are immutable objects. In the following, you are asked to write a code to "mutate" a string. In particular, you are given a string, a position number and a character. You should return a different string such that this string is the same as the input string except the character in the position is changed to the given character. 

Example:

    input: 
        original_string = "Hello world!"
        position=6
        character="W"
    ouput: "Hello World!"

```python
def problem2(original_string, position, character):
    return_string = ""
    
    for i in range( len( original_string ) ):
        
        if i == position:
            return_string = original_string.replace( original_string[i], character )
            
    return(return_string)
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 10, grade: true, solution: false -->
```python
# Test cases for Problem Set 4
assert problem2("Hello world!", 6, "W") == "Hello World!"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 10, grade: true, solution: false -->
```python
assert problem2("Hello world!", 11, ".") == "Hello world."
```

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
#Write ChatGPT code here]

```

## Problem 3 (Function)

Write a function that checks whether two words are anagrams. Two words are anagrams if they contain the same letters. For example, silent and listen are anagrams, dog and god are anagrams.

Your code must be able to deal with follwoing three cases:

Input: "silent", "listen"

output: silent and listen are anagrams

        1
        
Input: "dog", "good"

output: The two strings do not have the same lenght

        -1

Input: "cat", "tag"

output: cat and tag are not anagrams

        0

```python
def problem3(s1, s2):
    
    if len(s1) != len(s2):
        return -1
    else:
        sort1 = sorted( s1 )
        sort2 = sorted( s2 )
        if sort1 == sort2:
            return 1
        else:
            return 0
```

```python
assert problem3("silent", "listen") == 1
```

```python
assert problem3("dog", "good") == -1
```

```python
assert problem3("create", "regret") == 0
```

```python
problem3("silent", "listen")
```

Output:
```
1
```

```python
problem3("dog", "good")
```

Output:
```
-1
```

```python
problem3("create", "regret")
```

Output:
```
0
```

```python
problem3("cat", "act")
```

Output:
```
1
```

```python
#Write ChatGPT code here

```